# Notebook 1B - Mon/Tues
## Using constants and special functions, Documenting work in a Jupyter Notebook.
## Total Points: 11
#### PHYS 225 Intro to Computational Physics, Fall 2020

## Entering and reading scientific notation in Python
A rough estimate for the number of stars in the universe is $N = 1\times10^{21}$.
In Python you you would express the same value in scientific notation as `N = 1e21`.
Python also prints out numbers in using this same notation.

Always use the `e` notation to enter in scientific notation because it is more efficient, easier to read in code, and less prone to human error.

```python
# It is easy to enter and read scientific notation using the "e" notation
N = 1e21
print(N)
print(N/2)
```

Output:
```
1e+21
5e+20
```

```python
# Calculating the scientific notation makes the value harder to read, and can hide mistakes.
N = 1.0*10**21
print(N)

# a small typo drastically changes the result
N = 1*10*21
print(N)
```

Output:
```
1e+21
210
```

```python
# The "e" notation is flexible.
N = 0.5e21  # decimals smaller than 1 can use used before the "e"
print(N)

# values bigger than 10 can be used before the "e"
N = 833e9  # 833 billion dollars (estimated US Gov't federal deficit for 2018)
print(N)
```

Output:
```
5e+20
833000000000.0
```

### Task 1: Using scientific notation in Python (1 point)
* Create a variable `N_USpop` with a value of $3.26\times10^{8}$.
* Create a variable `diameter_flu_virus` with a value of $120 \textrm{ nm}$. Enter the value in meters (not nm)!

<!-- grade_id: cell-d42e4b5e7d6b0213, points: 0, grade: false, solution: true -->
```python
"""
Create the variables N_USpop and diameter_flu_vius and set them to the values specified above.
"""
### BEGIN SOLUTION
N_USpop = 3.26e8
diameter_flu_virus = 120e-9
### END SOLUTION
```

<!-- grade_id: cell-bd1de369be46c90c, points: 1, grade: true, solution: false -->
```python
"""Check the values """

# First check that variables have proper names 
# A NameError is produced if the variable name is not correct. 
print(N_USpop)
print(diameter_flu_virus)

# check the value of diameter_flu_virus
assert abs(diameter_flu_virus - (120e-9))/abs(120e-9) < 1e-4 # autograded with tolerance

# the test case for N_USpop is hidden
### BEGIN HIDDEN TESTS
assert abs(N_USpop - (3.26e8))/abs(3.26e8) < 1e-4 # autograded with tolerance
### END HIDDEN TESTS
```

Output:
```
326000000.0
1.2e-07
```

## Calculations with constants and special functions

You will need to use `import` statements to access very useful variables and functions such as  
* mathematical constants (e.g., $\pi$ or $e$)
* physical constants (e.g., $c$, $\hbar$)
* special functions (e.g., $\sin(x)$, $e^x$, $\log(x)$)

### Mathematical constants
Most mathematical constants can be accessed through the `numpy` (i.e., "numerical python") package , which requires importing values from a package. We will import from the `numpy` package, which will be used throughout the course. A more complete discussion of importing packages will come in week 3.

`from numpy import pi, e  # imports the variables pi and e`  

```python
from numpy import pi, e  # imports the variables pi and e
print(pi)
print(e)
print(2*pi)
```

Output:
```
3.141592653589793
2.718281828459045
6.283185307179586
```

### Special functions (e.g., $\sin(x)$, $e^x$, $\log(x)$)
A large number of special functions, especially those common to a handheld calculator, are available in `numpy`. Please check out the [full list of available math functions](https://docs.scipy.org/doc/numpy-1.13.0/reference/routines.math.html) in numpy. 

A slightly different `import` statement will be used, which gives access to ALL the functions and variables in `numpy`.

`import numpy as np  # gives access to all the the numpy functions`

```python
""" Try out sin(x) in python """

import numpy as np  # you only need to execute the import once in a notebook, usually near the top.

# the np. is required because sin() is part of the numpy package
x1 = np.sin(0)
x2 = np.sin(pi/4)  # you can still use the pi value defined above
x3 = np.sin(pi/2)
print(x1)
print(x2)
print(x3)
```

Output:
```
0.0
0.7071067811865475
1.0
```

```python
""" Try out e^x = exp(x) in python """

# the np. is required because exp() is part of the numpy package
x1 = np.exp(0)
x2 = np.exp(1)  # you can still use the pi value defined above
x3 = np.exp(-1)
print(x1)
print(x2)
print(x3)
```

Output:
```
1.0
2.718281828459045
0.36787944117144233
```

### Physical constants
A useful set of physical constants can be imported from the `scipy.constants` package. Please take a quick look at the [full set of physical constants available](https://docs.scipy.org/doc/scipy/reference/constants.html) in `scipy.constants`.

```python
from scipy.constants import hbar, c, e, m_p
print(c)  # speed of light
print(hbar)
print(e)  # note that the value of e (which was 2.718...) is replaced by the electron charge.
print(m_p)
```

Output:
```
299792458.0
1.0545718176461565e-34
1.602176634e-19
1.67262192369e-27
```

### Task 2: Calculation of the speed of light from permittivity and permeability of free space (1 point)
A remarkable equation derived from electricity and magnetism showed that the speed of light $c$ was directly related to fundmanetal constants in electricity and magnetism, specifically the permittivity of free space $\epsilon_0$ and the permiability of free space $\mu_0$.
$$ c = \frac{1}{\sqrt{\epsilon_0 \mu_0}} $$

* Calculate the value of speed of light from the two EM constants and store the value in a variable name `c_EM`. Print it out and compare to the value in scipy.constants
* Calculate the difference between the built-in speed of light `c` in `scipy.constants` vs `c_EM`, which you just calculated. Store the value in a variable named `diff`


<!-- grade_id: cell-8ce42bb50c298bdf, points: 0, grade: false, solution: true -->
```python
"""
Compute the speed of light from EM constants and store 
the value in a variable named c_EM. Print out c_EM.

Compute the difference 

"""

# import the necessary constants first

# then calculate c_EM

### BEGIN SOLUTION
from scipy.constants import epsilon_0, mu_0, c

c_EM  = (epsilon_0 * mu_0)**(-0.5)
print(c_EM)

diff = c - c_EM
print(diff)

### END SOLUTION
```

Output:
```
299792458.0000065
-6.496906280517578e-06
```

<!-- grade_id: cell-a2d1c2db0e5b32f5, points: 1, grade: true, solution: false -->
```python
"""Check that c_EM  and diff calculations"""
assert abs(c_EM - (299792458.0))/abs(299792458.0) < 1e-9 # autograded with tolerance
assert  abs(diff) < 1e-1 # autograded with tolerance
```

### Task 3: Calculation of the Rydberg energy (1 point)
The Rydberg energy, $R$, gives the energy scale for energy transmissions of a Hydrogen atom and can be derived from the Bohr model of the atom. The Rydberg energy, $R$, can be used to calculate the absorption or emission spectrum of Hydrogen. 
$$ R = \frac{m_e e^4}{8 h^2 \epsilon_0^2}  \approx 2.18\times10^{-18} \textrm{ J} = 13.6 \textrm{ eV}$$

<!-- grade_id: cell-1435d38a486eada5, points: 0, grade: false, solution: true -->
```python
"""
Compute the Rydberg Energy and save the value to a variable named R. Print out R.
Make sure you get all your constants from scipy.constants.
"""

# import the necessary constants first

# then calculate R

### BEGIN SOLUTION
from scipy.constants import h, epsilon_0, e, m_e
R = m_e * e**4 / (8 * h**2 * epsilon_0**2)
print(R)
print(R/e)
### END SOLUTION
```

Output:
```
2.1798723610862185e-18
13.605693122885842
```

<!-- grade_id: cell-e2485ddc8868d8bc, points: 1, grade: true, solution: false -->
```python
"""Check that R = 2.179872325390254e-18"""
assert abs(R - (2.179872325390254e-18))/abs(2.179872325390254e-18) < 1e-5 # autograded with tolerance
```

## Documenting your solutions in a Jupyter notebook 
As you saw on Monday/Tuesday, Jupyter notebooks can contain images, snazzy LaTeX equations, and formatted text, such as headings, bold, italics. Additional useful features are hyperlinks and lists.  

A full list of formatting features and examples of how to use it are given at the **[official markdown syntax website](https://daringfireball.net/projects/markdown/syntax)**.

## Pushing a broken down car

Dr. Zwickl's brother-in-law, Matt, was an offensive lineman on his college football team and is way-way stronger than Dr. Zwickl. On a recent trip back from a pizza restaurant, Matt was describing how he pushed his car up a hill in their neighborhood after the transmission stopped working. Dr. Zwickl was surprised Matt was able to push the car up the hill given his own troubles pushing a minivan on a flat driveway. 

However, it made him wonder, "What is the steepest hill I could push my car up?" 

![Minivan](2019_honda_odyssey_angularfront.jpg)

### Task 4: Get your pencil and paper solution into a Jupyter notebook (2 point)
On a separate sheet of paper, neatly draw a diagram and derive a formula for the force required to push a minivan up a hill at constant speed. Upload images of your solution and your partners solution.
* The images must be uploaded to the same directory as the Jupyter notebook. 
* The markdown syntax is `![Text descibing image](filename)`
* The `filename` needs to include the extention (e.g., .jpg, .png)
* You might want to resize the image before uploading (optional)

<!-- grade_id: cell-0be11b80020154c6, points: 2, grade: true, solution: true -->
**Write Markdown to this cell to integrate the image of your solution.**

### Task 5: Compute the numerical value of the force required to push (1 point)
In order to test your equation, use very unrealistic numbers for a minivan on a hill, assume 
* the mass of the minivan is $m_\textrm{minivan} = 10 \textrm{kg}$, and 
* the angle of the hill with respect to the horizontal is $\theta = 38^{\circ}$
* the van can roll with minimal friction because it is in neutral. 
* The value of $g = 9.8 \textrm{m/s}^2$.

Keep the value of the force in a variable named `force1` and print out `force1`.

<!-- grade_id: cell-85e6a66131a9c182, points: 0, grade: false, solution: true -->
```python
"""
Compute the numerical value for the above force and save the value to a variable named force1.
If numpy was imported earlier in the notebook (meaning the cell was run), 
then you don't need to import it again.
"""

# define your variables first

# then compute the force

### BEGIN SOLUTION
m = 10  # mass in kg
g = 9.8 # gravitational acceleration in m/s^2
theta = 38*pi/180 # angle in degrees converted to radians

force1 = m*g*np.sin(theta)
print(force1)
del  m, g, theta
### END SOLUTION
```

Output:
```
60.3348245819145
```

The value for the force should be about 60 N. A hidden autograder test will check the exact value.

<!-- grade_id: cell-ba9b58c5036a73f6, points: 1, grade: true, solution: false -->
```python
"""Check the value for the force"""
# check the variable name
print(force1)
### BEGIN HIDDEN TESTS
assert abs(force1 - (60.3348245819145))/abs(60.3348245819145) < 1e-2 # autograded with tolerance
### END HIDDEN TESTS
```

Output:
```
60.3348245819145
```

### Task 6: Select reasonable estimates for the minivan mass $m$ and your maximum push force $F_\textrm{max}$ (1 point)
* Describe how you determined your estimates. 
* Try to use physical reasoning to estimate a reasonable value for $F_\textrm{max}$ based on other feats of strength you can do.
* Google is great for finding parameters for common objects like minivans.

<!-- grade_id: cell-877ed77d6f82957c, points: 1, grade: true, solution: true -->


### Task 7: Estimate the steepest hill you could push a minivan up (3 points)
Show your code below. Because each group may use different parameters, the question is manually graded.
1. You must use the following variables for the initial parameter estimates and final angle:
    * `F_max`  (the max force you can push)
    * `M_minivan` (the mass of the minivan)
    * `theta_max` (the max you can push the minivan up)
1. You **must** add comments to each line (either directly above, or at the end of the line)
    * All variables should have a comment that specifies the quantity and units 
    * For estimated parameters, specify what your estimate is based on.
    * Answers with insufficient comments will lose 1 point.
1. Your final answer for `theta_max` should be in **degrees**.
1. Print the final calculated value of `theta_max`
1. Review your final answer with instructor before leaving.


<!-- grade_id: cell-3ab8395aabf71ea4, points: 0, grade: false, solution: true -->
```python
"""Estimate the steepest angle you could push a minivan."""
# define all your variables

# compute your angle

### BEGIN SOLUTION
# curb weight in pounds, coverted to kg, based on Honda website, for 2018
M_minivan = 4354*0.453592 

g = 9.8 # gravitational acceleration in m/s^2

# max push force in Pounds, converted to kg then newtons. Assumes I can push with 200 lbs
F_max = 200 * 0.453592 * 9.8


theta_max = np.arcsin( F_max/(M_minivan * g) ) * 180 / pi # angle in degrees converted to radians

print(theta_max)
### END SOLUTION
```

Output:
```
2.6327950260150352
```

<!-- grade_id: cell-af38ee5e3acdeb8e, points: 3, grade: true, solution: false -->
```python
"""Check the calculation of theta_max"""
# check the variable name
print(theta_max)
assert M_minivan > 0 # Minivan mass variable should exist and be positive
assert F_max > 0 # F_max (max push) variable should exist and be positive
assert theta_max > 0.1 # If you get a number below 0.1, check units on your angle. Radians or degrees?
### BEGIN HIDDEN TESTS
assert M_minivan > 3000*0.453592  # a light car is 3000 pounds
assert M_minivan < 6000*0.453592 # a heavy truck is 6000 pounds, the minivan is right in the middle.
assert F_max > 30 *0.453592 *9.8  # 50 pounds is like a heavy suitcase
assert F_max < 600 *0.453592 *9.8  # strongest powerlifters can't do more than this...crazy to suppose.
# smallest possible angle
assert theta_max > np.arcsin( 30 / 6000)  * 180 / pi  # 50 pounds push, 6000 pound car
assert theta_max < np.arcsin( 600/ 3000) * 180 / pi # 600 pounds push, 3000 pound car
print("smallest reasonable theta_max", np.arcsin( 30 / 6000)  * 180 / pi)
print("largest reasonable theta_max", np.arcsin( 600 / 3000)  * 180 / pi)
### END HIDDEN TESTS
```

Output:
```
2.6327950260150352
smallest reasonable theta_max 0.2864800912409137
largest reasonable theta_max 11.536959032815489
```

### Task 8: Use LaTeX to write these two equations from your solution (2 points).

1. An equation for the push force (calculated in task 5) in terms of masses and angles.
1. An equation for the angle (calculated in Task 7) in terms of mass and push force.

**Inline equations** are used to put math in the middle of a sentence. The `$` marks the beginning and end of the equation. In the Markdown it appears as `$E=mc^2$` but will print as $E=mc^2$.

**Centered equations** aappear on their own line and are marked by `$$`. The Markdown will be entered as   
`$$E = mc^2$$`  
but will appear as 
$$E = mc^2$$

There are many LaTeX guides online, but here is link to one [LaTeX cheat sheet](https://users.dickinson.edu/~richesod/latex/latexcheatsheet.pdf). You can also ask Google stuff like "how do I make fractions in latex".


<!-- grade_id: cell-abc8d759b255e9d9, points: 1, grade: true, solution: true -->
Write your two LaTeX equations in here. Don't forget the `$` or `$$` before and after!

## Remember to have Dr. Zwickl check over your notebook before submitting. 